```python
#!pip install pyLDAvis
```

# Sprint Challenge
## *Data Science Unit 4 Sprint 1*

After a week of Natural Language Processing, you've learned some cool new stuff: how to process text, how turn text into vectors, and how to model topics from documents. Apply your newly acquired skills to one of the most famous NLP datasets out there: [Yelp](https://www.yelp.com/dataset). As part of the job selection process, some of my friends have been asked to create analysis of this dataset, so I want to empower you to have a head start.  

The real dataset is massive (almost 8 gigs uncompressed). I've sampled the data for you to something more manageable for the Sprint Challenge. You can analyze the full dataset as a stretch goal or after the sprint challenge. As you work on the challenge, I suggest adding notes about your findings and things you want to analyze in the future.

## Challenge Objectives
Successfully complete all these objectives to earn full credit.

**Successful completion is defined as passing all the unit tests in each objective.**  

Each unit test that you pass is 1 point.

There are 5 total possible points in this sprint challenge.


There are more details on each objective further down in the notebook.*
* <a href="#p1">Part 1</a>: Write a function to tokenize the yelp reviews
* <a href="#p2">Part 2</a>: Create a vector representation of those tokens
* <a href="#p3">Part 3</a>: Use your tokens in a classification model on yelp rating
* <a href="#p4">Part 4</a>: Estimate & Interpret a topic model of the Yelp reviews

____

# Before you submit your notebook you must first

1) Restart your notebook's Kernel

2) Run all cells sequentially, from top to bottom, so that cell numbers are sequential numbers (i.e. 1,2,3,4,5...)
- Easiest way to do this is to click on the **Cell** tab at the top of your notebook and select **Run All** from the drop down menu.

3) Comment out the cell that generates a pyLDAvis visual in objective 4 (see instructions in that section).
____



### Import Data

<!-- grade_id: cell-395851cd95d17235, points: 0, grade: false, solution: true -->
```python
import pandas as pd

# Load reviews from URL
data_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_4/unit1_nlp/review_sample.json'

# Import data into a DataFrame named df
df = pd.read_json(data_url, lines=True)
df.shape
```

Output:
```
(10000, 9)
```

<!-- grade_id: cell-cb5006475e42b8f9, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert isinstance(df, pd.DataFrame), 'df is not a DataFrame. Did you import the data into df?'
assert df.shape[0] == 10000, 'DataFrame df has the wrong number of rows.'
```

## Part 1: Tokenize Function
<a id="#p1"></a>

Complete the function `tokenize`. Your function should
- accept one document at a time
- return a list of tokens

You are free to use any method you have learned this week.

```python
# Optional: Consider using spaCy in your function. The spaCy library can be imported by running this cell.
# A pre-trained model (en_core_web_sm) has been made available to you in the CodeGrade container.
# If you DON'T need use the en_core_web_sm model, you can comment it out below.
import spacy
nlp = spacy.load('en_core_web_sm')
```

Output:
```
/usr/local/lib/python3.8/dist-packages/torch/cuda/__init__.py:497: UserWarning: Can't initialize NVML
  warnings.warn("Can't initialize NVML")
```

<!-- grade_id: cell-3d570d5a1cd6cb64, points: 0, grade: false, solution: true -->
```python
def tokenize(doc):
  """
  Using SpaCy, we take in a doc and return a list of tokens as lemmas
  Removes stop words and punctuation
  """
  document = nlp(doc)

  return [token.lemma_.strip() for token in document if (token.is_stop != True) and (token.is_punct != True)]
```

```python
def tokenize(doc):
  lemma_list = [token.lemma_.lower().strip() for token in nlp(doc) if (not token.is_stop) and (not token.is_punct) and (not token.is_space)]
  return lemma_list
```

```python
tokenize(df.sample(n=1)["text"].iloc[0])
```

Output:
```
['love',
 'Grimaldi',
 'visit',
 'awful',
 'location',
 'server',
 'Dustin',
 'abhorrent',
 'rude',
 'order',
 'drink',
 'rude',
 'order',
 'food',
 'asshole',
 'get',
 'drink',
 'refill',
 'chore',
 'glass',
 'deliver',
 'second',
 'pizza',
 'say',
 'refill',
 'come',
 'manage',
 'drink',
 'refill',
 'table',
 'make',
 'eye',
 'contact',
 'come',
 'ready',
 'drop',
 'ticket',
 'point',
 'ask',
 'drink',
 'refill',
 'entire',
 'PIZZA',
 'FINISHED',
 'lean',
 'entire',
 'table',
 'refill',
 'shove',
 'crotch',
 'mother',
 'face',
 'way',
 'guess',
 'kill',
 'walk',
 'round',
 'table',
 'love',
 'pizza',
 'plenty',
 'location',
 'treat',
 'like',
 'want',
 'dad',
 'regular',
 'specific',
 'location',
 'favorite',
 'spot',
 'local',
 'high',
 'tolerance',
 'bad',
 'service',
 'bad',
 'experience',
 'height',
 'day',
 'restaurant',
 'opt',
 'dine',
 'lunch',
 'rush',
 'care',
 'Dustin',
 'have',
 'bad',
 'day',
 'able',
 'tell',
 'offer',
 'serve',
 'slice',
 'happen',
 'grimaldi',
 'offer
```

<!-- grade_id: cell-02da164f6fbe730a, points: 1, grade: true, solution: false -->
```python
'''Testing'''
assert isinstance(tokenize(df.sample(n=1)["text"].iloc[0]), list), "Make sure your tokenizer function accepts a single document and returns a list of tokens!"
```

## Part 2: Vector Representation
<a id="#p2"></a>
1. Create a vector representation of the reviews (i.e. create a doc-term matrix).
2. Write a fake review and query for the 10 most similar reviews, print the text of the reviews. Do you notice any patterns?
    - Given the size of the dataset, use `NearestNeighbors` model for this.

<!-- grade_id: cell-0e96491cb529202c, points: 0, grade: false, solution: true -->
```python
%%time
# Create a vector representation of the reviews
# Name that doc-term matrix "dtm"

from sklearn.feature_extraction.text import CountVectorizer

vect = CountVectorizer()
vect.fit(df['text'])
dtm = vect.transform(df['text'])
```

Output:
```
CPU times: user 1.84 s, sys: 87.6 ms, total: 1.92 s
Wall time: 2.29 s
```

```python
print(dtm)

type(dtm)
```

Output:
```
(0, 185)	1
  (0, 278)	1
  (0, 566)	1
  (0, 600)	1
  (0, 636)	1
  (0, 833)	1
  (0, 1227)	1
  (0, 1295)	1
  (0, 1357)	1
  (0, 1483)	1
  (0, 1500)	4
  (0, 2067)	1
  (0, 2598)	1
  (0, 2608)	2
  (0, 2881)	1
  (0, 2938)	1
  (0, 3889)	1
  (0, 3890)	1
  (0, 3977)	1
  (0, 4067)	1
  (0, 5508)	1
  (0, 5818)	1
  (0, 6080)	1
  (0, 9143)	3
  (0, 9523)	1
  :	:
  (9999, 25201)	1
  (9999, 25443)	1
  (9999, 25769)	1
  (9999, 25838)	1
  (9999, 25868)	1
  (9999, 26489)	10
  (9999, 26492)	1
  (9999, 26507)	1
  (9999, 26509)	1
  (9999, 26522)	4
  (9999, 26609)	1
  (9999, 26714)	3
  (9999, 26761)	1
  (9999, 26765)	1
  (9999, 26783)	1
  (9999, 26818)	2
  (9999, 26907)	2
  (9999, 26913)	1
  (9999, 26945)	1
  (9999, 26951)	1
  (9999, 27025)	1
  (9999, 27029)	1
  (9999, 27162)	1
  (9999, 27235)	2
  (9999, 27246)	1
scipy.sparse.csr.csr_matrix
```

<!-- grade_id: cell-3d5bc610a8ec6b24, points: 0, grade: false, solution: true -->
```python
# Create and fit a NearestNeighbors model named "nn"
from sklearn.neighbors import NearestNeighbors

nn = NearestNeighbors(n_neighbors=10, algorithm='kd_tree')
nn.fit(dtm)
```

Output:
```
/usr/local/lib/python3.8/dist-packages/sklearn/neighbors/_base.py:504: UserWarning: cannot use tree with sparse input: using brute force
  warnings.warn("cannot use tree with sparse input: using brute force")
NearestNeighbors(algorithm='kd_tree', n_neighbors=10)
```

<!-- grade_id: cell-c43704dcff67e99b, points: 1, grade: true, solution: false -->
```python
'''Testing.'''
assert nn.__module__ == 'sklearn.neighbors._unsupervised', ' nn is not a NearestNeighbors instance.'
assert nn.n_neighbors == 10, 'nn has the wrong value for n_neighbors'
```

<!-- grade_id: cell-496203e8746296ca, points: 0, grade: false, solution: true -->
```python
# Create a fake review and find the 10 most similar reviews

fake_review = ["""
Probably the best coffee place in SJ. Super friendly staff! Offers takeout coffee with pre-order during the pandemic, as well as pastry items and other things. I got a french press from them as my former one broke
""" ]
doc = vect.transform(fake_review)
neigh_dist, neigh_index = nn.kneighbors(doc)
neigh_index

```

Output:
```
array([[9981, 4569, 4730, 5983, 4626, 5129,  964, 6810, 7586, 8858]])
```

## Part 3: Classification
<a id="#p3"></a>
Your goal in this section will be to predict `stars` from the review dataset.

1. Create a pipeline object with a sklearn `CountVectorizer` or `TfidfVector` and any sklearn classifier.
    - Use that pipeline to train a model to predict the `stars` feature (i.e. the labels).
    - Use that Pipeline to predict a star rating for your fake review from Part 2.



2. Create a parameter dict including `one parameter for the vectorizer` and `one parameter for the model`.
    - Include 2 possible values for each parameter
    - **Use `n_jobs` = 1**
    - Due to limited computational resources on CodeGrader `DO NOT INCLUDE ADDITIONAL PARAMETERS OR VALUES PLEASE.`
    
    
3. Train the entire pipeline with a GridSearch
    - Name your GridSearch object as `gs`

<!-- grade_id: cell-e2beb0252d274bba, points: 0, grade: false, solution: true -->
```python
%%time
from sklearn.pipeline import Pipeline
from sklearn.feature_extraction.text import TfidfVectorizer
from sklearn.neighbors import KNeighborsClassifier
from sklearn.model_selection import GridSearchCV
from sklearn.ensemble import RandomForestClassifier

# Name the gridsearch instance "gs"
X = df.text
y = df.stars

tfidf= TfidfVectorizer(stop_words='english', tokenizer=None)

rfc = RandomForestClassifier(random_state=42)

pipe= Pipeline([('vect', tfidf),
                 ("clf", rfc)])

parameters = {
    'vect__max_df' : (0.75,1.0),
    'clf__max_depth' : (10,20)
}

gs = GridSearchCV(pipe,parameters, n_jobs=1)
gs.fit(X,y)

print(X.shape)
print(y.shape)

```

Output:
```
(10000,)
(10000,)
CPU times: user 40.7 s, sys: 118 ms, total: 40.8 s
Wall time: 40.9 s
```

<!-- grade_id: cell-d07134c6fe5d056e, points: 1, grade: true, solution: false -->
```python
# Visible Testing
prediction = gs.predict(["I wish dogs knew how to speak English."])[0]
assert prediction in df.stars.values, 'You gs object should be able to accept raw text within a list. Did you include a vectorizer in your pipeline?'
```

## Part 4: Topic Modeling

Let's find out what those yelp reviews are saying! :D

1. Estimate a LDA topic model of the review text
    - Set num_topics to `5`
    - Name your LDA model `lda`
2. Create 1-2 visualizations of the results
    - You can use the most important 3 words of a topic in relevant visualizations. Refer to yesterday's notebook to extract.
3. In markdown, write 1-2 paragraphs of analysis on the results of your topic model

When you instantiate your LDA model, it should look like this:

```python
lda = LdaModel(corpus=corpus,
               id2word=id2word,
               random_state=723812,
               num_topics = num_topics,
               passes=1
              )

```

__*Note*__: You can pass the DataFrame column of text reviews to gensim. You do not have to use a generator.

## Note about  pyLDAvis

**pyLDAvis** is the Topic modeling package that we used in class to visualize the topics that LDA generates for us.

You are welcomed to use pyLDAvis if you'd like for your visualization. However, **you MUST comment out the code that imports the package and the cell that generates the visualization before you submit your notebook to CodeGrade.**

Although you should leave the print out of the visualization for graders to see (i.e. comment out the cell after you run it to create the viz).

```python
from gensim import corpora
# Due to limited computationalresources on CodeGrader, use the non-multicore version of LDA
from gensim.models.ldamodel import LdaModel
import gensim
import re
```

### 1. Estimate a LDA topic model of the review tex

<!-- grade_id: cell-66331a185ff52f15, points: 0, grade: false, solution: true -->
```python
# Remember to read the LDA docs for more information on the various class attirbutes and methods available to you
# in the LDA model: https://radimrehurek.com/gensim/models/ldamodel.html

# don't change this value
num_topics = 5

# use tokenize function you created earlier to create tokens
tokens = df.text.apply(tokenize)
# create a id2word object (hint: use corpora.Dictionary)
id2word = corpora.Dictionary(tokens)
# create a corpus object (hint: id2word.doc2bow)
corpus = [id2word.doc2bow(text) for text in tokens]
# instantiate an lda model
lda = LdaModel(corpus = corpus,
               id2word = id2word,
               num_topics = num_topics,
               random_state = 42,
               passes = 1)
```

Output:
```
WARNING:gensim.models.ldamodel:too few updates, training might not converge; consider increasing the number of passes or iterations to improve accuracy
```

#### Testing

<!-- grade_id: cell-5a3c181311134fa9, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert lda.get_topics().shape[0] == 5, 'Did your model complete its training? Did you set num_topics to 5?'
```

#### 2. Create 1-2 visualizations of the results

<!-- grade_id: cell-9b043e992fbd218c, points: 0, grade: false, solution: true -->
```python
import seaborn as sns
import matplotlib.pyplot as plt
import pyLDAvis.gensim_models

# Use pyLDAvis (or a ploting tool of your choice) to visualize your results

# YOUR CODE HERE
#pyLDAvis.enable_notebook()
#vis = pyLDAvis.gensim_models.prepare(lda, corpus, id2word)
#vis
```

Output:
```
ModuleNotFoundError: No module named 'pyLDAvis'

---------------------------------------------------------------------------
NOTE: If your import is failing due to a missing package, you can
manually install dependencies using either !pip or !apt.

To view examples of installing some common dependencies, click the
"Open Examples" button below.
---------------------------------------------------------------------------
```

#### 3. In markdown, write 1-2 paragraphs of analysis on the results of your topic model

### Model Analysis:

Today, we used a Latent Dirichlet Allocation model. This model tokenized our text data and created a dictionary based off those words. That dictionary was then taken into a corpus using list comprehension and crafted into an LDA model. We can see that the model has determined our topics 1 and 4 have high correlation(or have words that are similar in meaning) topic 2 is close to these two as well, however, our other topics seem to be much more spread apart. This may also be because of the limit in topics and our relevance level.


The results of this model are subpar in my opinion. The text was not cleaned very well, there was a limit on how much we could tune the hyperparameters. This led to, what I would consider, a degredation in the true accuracy of the topics in our documents. I would be interested to see what we could accomplish by playing with the data and parameters.