## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.

### Instructions

* **Download this notebook** as you would any other ipynb file
* **Upload** to Google Colab or work locally (if you have that set-up)
* **Delete `raise NotImplementedError()`**
* Write your code in the `# YOUR CODE HERE` space
* **Execute** the Test cells that contain `assert` statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)
* **Save** your notebook when you are finished
* **Download** as a `ipynb` file (if working in Colab)
* **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)

# Unit 1 Sprint 2 Module 1

## Hypothesis Testing and Confidence Intervals

### Objectives

* Explain the purpose of a t-test and identify applications
* Use a t-test for independence to test for a statistically significant association between two categorical variables
* Use a t-test p-value to draw the correct conclusions about the null and alternative hypothesis
* Demonstrate how to build a confidence interval around a sample estimate

#### Total notebook points: 15

### Introduction

Mosquito nets have traditionally been an important tool to prevent mosquito bites in parts of the world where malaria is endemic. However, it may not be practical for an army that is on the move to set up and carry mosquito nets each night and day. Impregnating soldiers’ uniforms with insect repellent solves the mobility problem but also has drawbacks. First, the insect repellent quickly becomes ineffective with repeated washing and ironing and must be frequently reapplied. Second, in hot and humid climates the insect repellent can be absorbed through the skin, and the long-term effects of this exposure are unknown. One compromise is to have soldiers apply patches treated with insect repellent to their clothing. These patches would last longer because they would not be washed or ironed and would not expose the entire body to the insect repellent.

### Dataset Description

The `Mosquito.xlsx` dataset contains data recorded in an experiment conducted on male soldiers in the Indian Army who were stationed in the Tezpur/Solmara garrison in Northeast India. Thirty soldiers were randomly selected to receive one of three types of mosquito single repellent patch. After giving informed consent, the study participants affixed the patches at predetermined points on their uniforms and research assistants (who were blinded to the type of repellent used) counted the number of times a mosquito landed on each individual in an hour. 

Medical officers with the Indian Army have recorded data on mosquito bites and related illness for many years and can say with authority that the mean number of mosquito touches for soldiers not wearing any mosquito repellent is 8.2 per hour. **We wish to determine if wearing a single repellent patch changes the mean number of mosquito touches for soldiers compared to not wearing any mosquito repellent.**

*Adapted from: A. Bhatnagar and V.K. Mehta (2007). "Efficacy of Deltamethrin and Cyfluthrin Impregnated Cloth Over Uniform Against Mosquito Bites," Medical Journal Armed Forces India, Vol. 63, pp. 120-122.*

**Task 1** - Load the data

Let's load the data! The URL has been provided as well as the imports for pandas and numpy.

* load your CSV file into a DataFrame named `df_mosquito`

<!-- grade_id: cell-dae3789e00dd9b29, points: 0, grade: false, solution: true -->
```python
# Task 1
import pandas as pd
import numpy as np

# URL for the dataset
data_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Mosquito/Mosquito.csv'

### BEGIN SOLUTION
df_mosquito = pd.read_csv(data_url)
### END SOLUTION

# Print out your DataFrame
df_mosquito.head()
```

Output:
```
ID  Mosq_count
0   1           4
1   2          10
2   3          13
3   4           0
4   5          11
```

**Task 1 Test**

<!-- grade_id: cell-a347d8dd7572e6d4, points: 1, grade: true, solution: false -->
```python
# Task 1 - Test

assert isinstance(df_mosquito, pd.DataFrame), 'Have you created a DataFrame named `df_mosquito`?'
assert len(df_mosquito) == 90

### BEGIN HIDDEN TESTS
assert df_mosquito.shape == (90, 2), 'Make sure you correctly read in the DataFrame'
### END HIDDEN TESTS
```

**Task 2** - Calculate the mean

* Calculate the mean number of mosquito touches in the sample. Assign your answer to the variable `mosquito_touch_mean`.

<!-- grade_id: cell-22ced5cd2cae6230, points: 0, grade: false, solution: true -->
```python
# Task 2

### BEGIN SOLUTION
mosquito_touch_mean = df_mosquito['Mosq_count'].mean()
### END SOLUTION
```

**Task 2 Test**

<!-- grade_id: cell-3970fbb1a3eaefb5, points: 1, grade: true, solution: false -->
```python
# Task 2 - Test
# Hidden tests - you will see the results when you submit to Canvas
### BEGIN HIDDEN TESTS
assert round(mosquito_touch_mean) == 8, 'Did you use the `.mean()` method?'
### END HIDDEN TESTS
```

**Task 3** - Calculate the standard deviation

* Calculate the standard deviation of the number of mosquito touches in the sample. Assign your answer to `mosquito_touch_std`.

<!-- grade_id: cell-9c4fe58a9fe2d7aa, points: 0, grade: false, solution: true -->
```python
# Task 3

### BEGIN SOLUTION
mosquito_touch_std = df_mosquito['Mosq_count'].std()
### END SOLUTION
```

**Task 3 Test**

<!-- grade_id: cell-4939c4289c39b4e3, points: 1, grade: true, solution: false -->
```python
# Task 3 - Test
# Hidden tests - you will see the results when you submit to Canvas
### BEGIN HIDDEN TESTS
assert round(mosquito_touch_std) == 3, 'Did you use the .std() method?'
### END HIDDEN TESTS
```

**Task 4 -** Statistical hypotheses

From the list of choices below, select the null and alternative hypotheses using the experiment information described above.  Specify your answer in the next code block using `Answer = `.  For example, if the correct answer is choice B, you'll type `Answer = 'B'`.

A: $H_0: \mu = 8$ vs. $H_a: \mu = 8$

B: $H_0: \mu \neq 8.2$ vs. $H_a: \mu = 8$

C: $H_0: \mu \neq 8.2$ vs. $H_a: \mu = 8.2$ 

D: $H_0: \mu =8.2$ vs. $H_a: \mu \neq 8.2$ 

<!-- grade_id: cell-5ab063cfe94f3929, points: 0, grade: false, solution: true -->
```python
# Task 4

### BEGIN SOLUTION
Answer = 'D'
### END SOLUTION

```

**Task 4 Test**

<!-- grade_id: cell-085e920bd147cbab, points: 1, grade: true, solution: false -->
```python
# Task 4 - Test
# Hidden tests - you will see the results when you submit to Canvas
### BEGIN HIDDEN TESTS
assert Answer == 'D', 'Think carefully about your hypotheses.'
### END HIDDEN TESTS
```

**Task 5** - One-sample t-test

* Conduct a 1-sample t-test to test your hypotheses. Assign your t-test result to the variable `mosquito_pval`. **This should be a single value**.

*Hint: The `stats.ttest_1samp()` function returns two values; assign the results of the t-test to `_, mosquito_pval`.*

<!-- grade_id: cell-44b3f2a39081c5df, points: 0, grade: false, solution: true -->
```python
# Task 5

# Use the 'ttest_1samp' from the stats package
from scipy import stats

### BEGIN SOLUTION
_, mosquito_pval = stats.ttest_1samp(df_mosquito['Mosq_count'], 8.2)
### END SOLUTION
```

**Task 5 Test**

<!-- grade_id: cell-3be2eb1f16e366e6, points: 1, grade: true, solution: false -->
```python
# Task 5 - Tests
# Hidden tests - you will see the results when you submit to Canvas
### BEGIN HIDDEN TESTS
assert round(mosquito_pval, 3) == 0.586, 'Did you use the correct population mean?'
### END HIDDEN TESTS
```

**Task 6**

Select the correct conclusion at the 0.05 significance level from the list of choices below. Specify your answer in the next code block using `Answer = `.  For example, if the correct answer is choice B, you'll type `Answer = 'B'`.

A: We reject the null hypothesis at the 0.05 significance level and can conclude that a single repellant patch reduces the mean number of mosquito touches.

B: We fail to reject the null hypothesis at the 0.05 significance level and can conclude that a single repellant patch reduces the mean number of mosquito touches.

C: We fail to reject the null hypothesis at the 0.05 significance level and can conclude that a single repellant patch does not change the mean number of mosquito touches.

D: We reject the null hypothesis at the 0.05 significance level and conclude that a single repellant patch does not increase the mean number of mosquito touches.


<!-- grade_id: cell-4d35984ef163c5e1, points: 0, grade: false, solution: true -->
```python
# Task 6

### BEGIN SOLUTION
Answer = 'C'
### END SOLUTION

```

**Task 6 Test**

<!-- grade_id: cell-2f182809a44a91f5, points: 1, grade: true, solution: false -->
```python
# Task 6 - Test
# Hidden tests - you will see the results when you submit to Canvas
### BEGIN HIDDEN TESTS
assert Answer == 'C', 'Think carefully about your conclusion.'
### END HIDDEN TESTS
```

**Task 7** - Calculate t*

For this task you will calculate t* for a 95% confidence interval.

* set the variable `deg_free` equal to the degrees of freedom for the `df_moquito` variable
* set the variable `t_star` equal to t* using `t.ppf(q, df)` with `q=0.975` and `df = deg_free`

Note: Don't worry about where the 0.975 value comes from - it relates to finding the *middle* of the 95% t-distribution. We're going to learn how to calculate the 95% confidence interval an easier way in the next exercise.

<!-- grade_id: cell-ab557b231ecefc1d, points: 0, grade: false, solution: true -->
```python
# Task 7

from scipy.stats import t

### BEGIN SOLUTION
deg_free = len(df_mosquito)-1
t_star = t.ppf(q=0.975, df=deg_free)
### END SOLUTION

# View your answer
print('t_star =', t_star)
```

Output:
```
t_star = 1.9869786993737677
```

**Task 7 - Test**

<!-- grade_id: cell-ea70eb8f030c2f57, points: 1, grade: true, solution: false -->
```python
# Task 7 Test
# Hidden tests - you will see the results when you submit to Canvas
### BEGIN HIDDEN TESTS
assert deg_free == 89, 'Remember the degrees of freedom is equal to n-1'
assert t_star == t.ppf(q=0.975, df=deg_free), 'Did you use the correct values for q and the degrees of freedom?';
### END HIDDEN TESTS
```

**Task 8** - Margin of error

In this task you'll calculate the margin of error for a 95% confidence interval (CI) for the mean number of mosquito touches.

* Assign the margin of error for a 95% CI to the variable `margin_err`

Hint: You already have the value for t* for a 95% CI and the standard error

<!-- grade_id: cell-fe91f19c7d908552, points: 0, grade: false, solution: true -->
```python
# Task 8

### BEGIN SOLUTION
margin_err = t_star * df_mosquito['Mosq_count'].sem()
### END SOLUTION

# View your answer
print('Margin of error = ', margin_err)
```

Output:
```
Margin of error =  0.6875174745924729
```

**Task 8 - Test**

<!-- grade_id: cell-2258dd78f578b1be, points: 1, grade: true, solution: false -->
```python
# Task 8 Test
# Hidden tests - you will see the results when you submit to Canvas
### BEGIN HIDDEN TESTS
assert deg_free == 89, 'Remember the degrees of freedom is equal to n-1'
assert t_star == t.ppf(q=0.975, df=89)
assert round(margin_err, 2) == .69
### END HIDDEN TESTS
```

**Task 9** - Calculate a confidence interval

For this task, you are going to calculate a 95% CI for the number of mosquito touches with the CI formula using the summary statistics and t* that you calculated above.

* Calculate the lower confidence level and assign it to `lower_cl`
* Calculate the upper confidence level and assign it to `upper_cl`

<!-- grade_id: cell-2ee7b7b25af4e5fc, points: 0, grade: false, solution: true -->
```python
# Task 9

### BEGIN SOLUTION
lower_cl = mosquito_touch_mean - margin_err
upper_cl = mosquito_touch_mean + margin_err
### END SOLUTION

# View your answers
print ('Lower confidence limit =', lower_cl)
print ('Upper confidence limit =', upper_cl)
```

Output:
```
Lower confidence limit = 7.3235936365186385
Upper confidence limit = 8.698628585703585
```

**Task 9 - Test**

<!-- grade_id: cell-8ea440c5365e886d, points: 1, grade: true, solution: false -->
```python
# Task 9 Test
# Hidden tests - you will see the results when you submit to Canvas
### BEGIN HIDDEN TESTS
assert lower_cl == mosquito_touch_mean - margin_err, 'The lower CL is the mean minus the margin of error'
assert upper_cl == mosquito_touch_mean + margin_err, 'The upper CL is the mean plus the margin of error'
### END HIDDEN TESTS
```

**Task 9** - ANSWER

Write out the confidence interval you just calculated. Use the following format:

*Example: We are 95% confident that the true mean of the mosquito touches is between [lower CL, upper CL]*

This task will not be autograded - but it is part of completing the project.

YOUR ANSWER

**Task 10** - 95% confidence interval using t-interval

As promised in Task 4, we're going to calculate the confidence interval the easy way. We'll use the `t.interval()` function to calculate the 95% confidence interval.

* Assign the confidence interval to `t_int_95`
* `alpha` should be set equal to the confidence level
* `df` is the degrees of freedom
* `loc` is the sample mean
* `scale` is the standard error of the distribution

<!-- grade_id: cell-305db71f92847f07, points: 0, grade: false, solution: true -->
```python
# Task 10

### BEGIN SOLUTION
t_int_95 = t.interval(alpha=0.95, df=89, loc=mosquito_touch_mean, scale=df_mosquito['Mosq_count'].sem())
### END SOLUTION

# View your answer
print(t_int_95)
```

Output:
```
(7.3235936365186385, 8.698628585703585)
```

**Task 10 - Test**

<!-- grade_id: cell-752d977ae5d519c5, points: 1, grade: true, solution: false -->
```python
# Task 10 Test

# Hidden tests - you will see the results when you submit to Canvas
### BEGIN HIDDEN TESTS
assert t_int_95 == t.interval(alpha= 0.95, df=89, 
                                loc=mosquito_touch_mean, 
                                scale=df_mosquito['Mosq_count'].sem()), 'Double check your parameters'
### END HIDDEN TESTS
```

```python

```

## Use the following information to complete Tasks 11-15

### Introduction

More than 14,000 people finished the 2020 Disney Marathon held on January 12. The results by age and gender group are included in the `Disney.csv dataset`. 

**We wish to determine if the mean finishing time for male and female marathon runners is the same or if there is a difference in the mean finishing time between male and female marathon runners.**


[Source: Track Shack. 2020 Disney Marathon Race Results](https://www.trackshackresults.com/disneysports/results/wdw/wdw20/mar_results.php)

**Task 11** - Load the next dataset

Let's load the data! The URL has been provided.

* load your CSV file into a DataFrame named `df_disney`

<!-- grade_id: cell-2b389fa658a51cd2, points: 0, grade: false, solution: true -->
```python
# Task 11
import pandas as pd
# URL for Disney marathon dataset
data_url2 = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Disney_Marathon/Disney.csv'

### BEGIN SOLUTION
df_disney = pd.read_csv(data_url2)
### END SOLUTION
```

**Task 11 Test**

<!-- grade_id: cell-c1f16b75c543e0d1, points: 1, grade: true, solution: false -->
```python
# Task 11 - Test

assert isinstance(df_disney, pd.DataFrame), 'Have you created a DataFrame named `df_disney`?'
assert len(df_disney) == 14106

### BEGIN HIDDEN TESTS
assert df_disney.shape == (14106, 5), 'Make sure you correctly read in the DataFrame'
### END HIDDEN TESTS
```

**Task 12 -** Statistical hypotheses

From the list of choices below, select the null and alternative hypotheses using the experiment information described above.  Let $\mu_1$ be the mean finishing time for all male runners and $\mu_2$ be the mean finishing time for all female runners.

Specify your answer in the next code block using `Answer = `.  For example, if the correct answer is choice B, you'll type `Answer = 'B'`.

A: $H_0: \mu_1 \neq \mu_2$ vs. $H_a: \mu_1 = \mu_2$

B: $H_0: \mu_1 = \mu_2$ vs. $H_a: \mu_1 \neq \mu_2$

C: $H_0: \mu_1 > \mu_2$ vs. $H_a: \mu_1 < \mu_2$ 

D: $H_0: \mu_1 <  \mu_2$ vs. $H_a: \mu_1 > \mu_2$ 

<!-- grade_id: cell-de36ef90faad17d8, points: 0, grade: false, solution: true -->
```python
# Task 12

### BEGIN SOLUTION
Answer = 'B'
### END SOLUTION

```

**Task 12 Test**

<!-- grade_id: cell-492e80f1f416b2c9, points: 1, grade: true, solution: false -->
```python
# Task 12 - Test
# Hidden tests - you will see the results when you submit to Canvas
### BEGIN HIDDEN TESTS
assert Answer == 'B', 'Think carefully about your hypotheses.'
### END HIDDEN TESTS
```

**Task 13** - Create new series from a DataFrame

Create **Series** (a pandas DataFrame column is a Series):

* one containing finishing times for male participants (`male_finish`)
* one containing finishing times for female participants (`female_finish`)

*Hint: Check the size of your resulting Series - it should have only one column!*

<!-- grade_id: cell-4290d3dd2f668a25, points: 0, grade: false, solution: true -->
```python
# Task 13

### BEGIN SOLUTION
male_finish = df_disney['time'].loc[df_disney['gender'] == 'M']
female_finish = df_disney['time'].loc[df_disney['gender'] == 'F']
### END SOLUTION
```

**Task 13 Test**

<!-- grade_id: cell-5384320ccd7f5048, points: 1, grade: true, solution: false -->
```python
# Task 13 - Test
# Visible testing - use this to check your results!
assert male_finish.shape == (6577,), 'Make sure you selected M and only have a single column.'
assert female_finish.shape == (7529,), 'Make sure you selected F and only have a single column'

# NO hidden tests
```

**Task 14** - Conduct a 2-sample t-test

Conduct a 2-sample t-test to test your hypotheses:

* Assign the t-statistic to a variable called `disney_tval`
* Assign the p-value to a variable called `disney_pval`

**Note:** The function returns two values and you can assign them with one line (example):

`variable1, variable2` = `some.function(arguments)`

<!-- grade_id: cell-d391e7e149994d87, points: 0, grade: false, solution: true -->
```python
# Task 14

### BEGIN SOLUTION
from scipy import stats

disney_tval = stats.ttest_ind(male_finish, female_finish)[0]
disney_pval = stats.ttest_ind(male_finish, female_finish)[1]

disney_tval, disney_pval = stats.ttest_ind(female_finish, male_finish)
stats.ttest_ind(female_finish, male_finish)
# Alternate solution:

#stats.ttest_ind_from_stats(male_finish_mean, male_finish_std, len(male_finish), 
#                                 female_finish_mean, female_finish_std, len(female_finish), equal_var=True)
### END SOLUTION
```

Output:
```
Ttest_indResult(statistic=29.27857393997243, pvalue=5.485138013952879e-183)
```

**Task 14 Test**

<!-- grade_id: cell-35a1e3947ac6388a, points: 1, grade: true, solution: false -->
```python
# Task 14 - Test
# Hidden tests - you will see the results when you submit to Canvas
### BEGIN HIDDEN TESTS
assert np.abs(round(disney_tval)) == 29, 'Did you assign the correct values?'
assert round(disney_pval) == 0, 'Did you assign the correct values?'
### END HIDDEN TESTS
```

**Task 15**

Select the correct conclusion at the 0.05 significance level from the list of choices below. Specify your answer in the next code block using `Answer = `.  For example, if the correct answer is choice B, you'll type `Answer = 'B'`.

A: We reject the null hypothesis at the 0.05 significance level and conclude the mean finishing time for male and female marathon runners is different.

B: We fail to reject the null hypothesis at the 0.05 significance level and conclude the mean finishing time for male and female marathon runners is different.

C: We reject the null hypothesis at the 0.05 significance level and conclude the mean finishing time for male and female marathon runners is the same.

D: We fail to reject the null hypothesis at the 0.05 significance level and conclude the mean finishing time for male and female marathon runners is the same.


<!-- grade_id: cell-66e19643f07fb7af, points: 0, grade: false, solution: true -->
```python
# Task 15

### BEGIN SOLUTION
Answer = 'A'
### END SOLUTION

```

**Task 15 Test**

<!-- grade_id: cell-d166fbea62c826e8, points: 1, grade: true, solution: false -->
```python
# Task 15 - Test
# Hidden tests - you will see the results when you submit to Canvas
### BEGIN HIDDEN TESTS
assert Answer == 'A', 'Think carefully about your conclusion.'
### END HIDDEN TESTS
```

```python

```