## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.

### Instructions

* **Download this notebook** as you would any other ipynb file
* **Upload** to Google Colab or work locally (if you have that set-up)
* **Delete `raise NotImplementedError()`**
* Write your code in the `# YOUR CODE HERE` space
* **Execute** the Test cells that contain `assert` statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)
* **Save** your notebook when you are finished
* **Download** as a `ipynb` file (if working in Colab)
* **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)

# Unit 1 Sprint 1 Module 1

## Introduction to Python

### Objectives

* Create and work with variable data types
* Understand indexing and methods
* Practice dictionary indexing

#### Total notebook points: 8

## Autograded Module Projects

Welcome to the first Module Project of Unit 1! You will complete a project (sometimes also referred to as an assignment) after the Guided Project for each Module. There will be four Module Projects per Sprint and each project is designed to provide you with an opportunity to practice what you learned in the Canvas Warm-up material and the Guided Project with your instructors.

Throughout Unit 1, the Module Projects and the Sprint Challenges are *autograded*. You will complete your work in a Jupyter/Python notebook (the files that end with `.ipynb`) and then upload your completed notebook to Canvas and submit for grading. This autograding process will check your answers and provide more information about the errors in your notebook. You'll receive a score when the testing is complete.

So, let's get started! If you are reading this notebook, then you have the correct autograded version.

**Task 1** - Write Your Name

Create a variable called `your_name` that includes your name. 

<!-- grade_id: cell-6027a63ecafcefab, points: 0, grade: false, solution: true -->
```python
#YOUR CODE HERE
your_name = "Julian"


#Output or print the results of your variable. 
```

**Task 1 Test**

<!-- grade_id: cell-a6c4f852066437b1, points: 1, grade: true, solution: false -->
```python
#Task 1 - Test
assert isinstance(your_name, str)
print("Hello", your_name )


###BEGIN HIDDEN TESTS
assert isinstance(your_name, str)
###END HIDDEN TESTS

```

Output:
```
Hello Julian
```

**Task 2** - Variable data types

* Assign the number `10` to a variable called `first_num`
* Assign the number `5` to a variable called `second_num`
* Assign the number `2` to a variable called `num_two` 

<!-- grade_id: cell-e96ddb74e7505094, points: 0, grade: false, solution: true -->
```python
#Task 2 

#YOUR CODE HERE
first_num = 10
second_num = 5
num_two = 2


```

**Task 2 Test**

<!-- grade_id: cell-a33f13464286f445, points: 1, grade: true, solution: false -->
```python
#Task 2 - Test
# Hidden tests - you will see the results when you submit to Canvas
###BEGIN HIDDEN TESTS
assert second_num == 5, "Did you assign the values correctly?"
###END HIDDEN TESTS
```

**Task 3** - Operators
* Return the product of all three variables created in the task above. Assign it to an object called `product_total`.
* Return the floor division of `product_total` by 7. Assign it **as an interger** to an object called `floor_value`. 

<!-- grade_id: cell-311e511f461035fc, points: 0, grade: false, solution: true -->
```python

#Task 3

#YOUR CODE HERE
product_total = first_num*second_num*num_two
display(product_total)

floor_value = int(product_total/7)
floor_value
```

Output:
```
100
14
```

**Task 3 Test**

<!-- grade_id: cell-c96d884a0d9c28fe, points: 1, grade: true, solution: false -->
```python
#Task 3 Test

assert product_total == 100, "Check your operators."

# Hidden tests - you will see the results when you submit to Canvas
###BEGIN HIDDEN TESTS
assert floor_value == 14, "Did you change your value to an interger?"
###END HIDDEN TESTS

```

**Task 4a** Python Data Types
* Create a list of five objects: `red`, `green`, `blue`, `yellow`, `purple`. Assign this list **in alphabetical order** to an object called `five_colors`.  


<!-- grade_id: cell-32200b8130c2bd81, points: 0, grade: false, solution: true -->
```python

#Task 4a

#YOUR CODE HERE
five_colors = ['red', 'green', 'blue', 'yellow', 'purple']
five_colors.sort()
five_colors


```

Output:
```
['blue', 'green', 'purple', 'red', 'yellow']
```

**Task 4a Test**

<!-- grade_id: cell-041962ba22262d5e, points: 0.5, grade: true, solution: false -->
```python
#Task 4a Test
assert len(five_colors) == 5
# Hidden tests - you will see the results when you submit to Canvas
###BEGIN HIDDEN TESTS
assert five_colors[-1] == 'yellow'
###END HIDDEN TESTS
```

**Task 4b** Python Data Types
* Remove the value that is in the middle of the `five_colors` list. Assign this new list to an object called `four_colors`.

<!-- grade_id: cell-de568d701f9137ca, points: 0, grade: false, solution: true -->
```python
#Task 4b

#YOUR CODE HERE
five_colors.remove(five_colors[(len(five_colors)//2)])
four_colors = five_colors
four_colors
```

Output:
```
['blue', 'green', 'red', 'yellow']
```

**Task 4b Test**

<!-- grade_id: cell-9ab8dc3c20c2a9c1, points: 0.5, grade: true, solution: false -->
```python
#Task 4b Test
assert len(four_colors) == 4, "Try displaying or printing the output of your object to verify your results"
# Hidden tests - you will see the results when you submit to Canvas
###BEGIN HIDDEN TESTS
assert five_colors[1] == 'green', "Try rerunning your previous cells to reset your variables."
###END HIDDEN TESTS

```

**Task 5** - Indexing
* Use indexing to return the *third* item in the `four_colors` list. Assign this to an object called `third_color`.
* Next, use string indexing to return the first and the last letter of the `third_color` object. Assign it to an object called `first_and_last`.   


<!-- grade_id: cell-9809eb2fc05f4fa5, points: 0, grade: false, solution: true -->
```python
#Task 5

#YOUR CODE HERE
third_color = four_colors[2]
third_color

#okay answer
third_color[0::2]
#best answer
first_and_last = third_color[0] + third_color[-1]
first_and_last
```

Output:
```
'rd'
```

**Task 5 Test**

<!-- grade_id: cell-229ed9f4a0babeba, points: 1, grade: true, solution: false -->
```python
#Task 5 Test
assert third_color == 'red'
# Hidden tests - you will see the results when you submit to Canvas
###BEGIN HIDDEN TESTS
assert first_and_last == 'rd'
###END HIDDEN TESTS
```

**Task 6a** - Indexing
* Create a dictionary that includes the following key: value pairs: `charizard: fire, squirtle: water, bulbasaur: leaf`. Assign it to an object called `poke_dictionary` 

<!-- grade_id: cell-24b1a6230e320ec9, points: 0, grade: false, solution: true -->
```python
#Task 6a

#YOUR CODE HERE
poke_dictionary = {'charizard':'fire', 
                  'squirtle': 'water',
                  'bulbasaur': 'leaf'}
```

**Task 6a Test**

<!-- grade_id: cell-5ad0aaf04fb6881b, points: 0.5, grade: true, solution: false -->
```python
#Task 6a Test
assert len(poke_dictionary.keys()) == 3, "Verify your syntax. Did you create a dictionary object?"

# Hidden tests - you will see the results when you submit to Canvas
###BEGIN HIDDEN TESTS
assert poke_dictionary['charizard'] == 'fire'
###END HIDDEN TESTS

```

**Task 6b** - Indexing
* Create a new dictionary called `more_poke` that includes the following: `pikachu: electric, pidgey: air` 
* Combine the key and values of the `more_poke` with the original `poke_dictionary`.
>Hint: You may need to research how to combine two dictionaries. 
* Assign all the `dictionary values` of the `poke_dictionary` to an object called `poke_values`. 
>Hint: Your output for `poke_values` will look like: `dict_values(['fire', 'water', 'leaf', 'electric', 'air'])`

<!-- grade_id: cell-0b4582e6a68091ef, points: 0, grade: false, solution: true -->
```python
#Task 6b

#YOUR CODE HERE
more_poke = {'pikachu':'electric', 'pidgey':'air'}
poke_dictionary.update(more_poke)
poke_values= poke_dictionary.values()
poke_values
```

Output:
```
dict_values(['fire', 'water', 'leaf', 'electric', 'air'])
```

**Task 6b Tests**

<!-- grade_id: cell-c561f60e90446fd7, points: 0.5, grade: true, solution: false -->
```python
#Task 6b Test
assert isinstance(more_poke, dict), "Check your syntax. Did you create a dictionary object?"
assert len(poke_dictionary) == 5, "Did you properly combine both dictionaries?"
# Hidden tests - you will see the results when you submit to Canvas


###BEGIN HIDDEN TESTS
assert 'fire' in poke_values, "Did you include all the key/value pairs?"
assert 'air' in poke_values, "Did you include all the key/value pairs?"
###END HIDDEN TESTS

```

**Task 7** Loops
* Use a for loop to go through the `poke_dictionary` object. Return all the dictionary **values** as uppercase letters and the **keys** as lowercase values. 
* *Hint*: This [StackOverflow](https://stackoverflow.com/questions/50004310/convert-python-dictionary-to-uppercase) page offers a solution (with a slight adjustment), but there are multiple ways you can do this!

<!-- grade_id: cell-0535b02af29e01e7, points: 0, grade: false, solution: true -->
```python
#Task 7

#YOUR CODE HERE
for i in poke_dictionary:
    poke_dictionary[i] = poke_dictionary[i].upper()
    
#another answer
#poke_dictionary = {k:v.upper() for k,v in poke_dictionary.items()}
```

**Task 7 Test**

<!-- grade_id: cell-0e4a39d38c38f186, points: 1, grade: true, solution: false -->
```python
#Task 8 Test
assert poke_dictionary['pikachu'].isupper(), "Your dictionary values should all be uppercase"
###BEGIN HIDDEN TESTS
assert list(poke_dictionary.keys())[0].islower(), "Your dictionary keys should still be lowercase"
###END HIDDEN TESTS
```

**Task 8** Imports
* Import the `seaborn` and `pandas` library. Alias the libraries as `sns` and `pd` respectively. 
* Use the `sns` library to load a dataset of your choice. Assign it to an object called `df`. 
* Run the following code: `df.loc[0]`. What data type is this? Assign your answer to an object called `first_row_type`.


<!-- grade_id: cell-6ed26cd55b737be1, points: 0, grade: false, solution: true -->
```python
#Task 8

#YOUR CODE HERE
import seaborn as sns
import pandas as pd


#sns.get_dataset_names()
df = sns.load_dataset('taxis')
first_row_type = type(df.loc[0])
```

**Task 8 Test**

<!-- grade_id: cell-39cc971493c3004e, points: 1, grade: true, solution: false -->
```python
#Task 8 Test

assert isinstance(df, pd.DataFrame), "Did you import the pandas library? Did you check the datatype of the df object?"

# Hidden tests - you will see the results when you submit to Canvas
###BEGIN HIDDEN TESTS
assert first_row_type == pd.Series
###END HIDDEN TESTS
```